# BLU05 - Exercises Notebook

```python
import hashlib # for grading purposes
import math
import numpy as np
import pandas as pd
import requests
import sqlalchemy

from bs4 import BeautifulSoup
```

## Part A - SQL exercises

### Querying the MoviesDb with a SQL client

Open your favorite SQL client and connect to the MoviesDb.
The connection settings are the following.

* host: db.lisbondatascience.org
* port: 5432
* user: ldsa_student
* database: datawrangling
* schema: blu05
* password: XXX (shared through slack)

This is the database that we used in the learning notebooks. Just to help you remember, these are the tables in the database.

1. movie: has information about the movies in the database
2. actor: has information about the actors in the database
3. genre: has information about the movie genres in the database
4. oscar: has a list of movies that won an oscar, per year
5. movie_actor: has the relationships between movies and actors (as in, which actors appeared in which movies and characters did the actors play)
6. movie_genre: has the relationships between movies and genres (as in, which genres does each movie belong to)

You can preview these tables in the SQL client.

### Q1. Select the original title of the movie with imdb_id tt0114814 

Write a query that selects the original title of the movie with imbdb_id tt0114814, and run it in the SQL client.

Then, assign the result to variable tt0114814_original_title (just copy and paste the value).

<!-- grade_id: cell-03b22b93be7b587e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-92e8f002ca863db1, points: 1, grade: true, solution: false -->
```python
expected_hash = '5b747e07ed1745562a23e71038c445a3fd55c917c55b35cc8294b93735f12698'
assert hashlib.sha256(tt0114814_original_title.encode()).hexdigest() == expected_hash
```

### Q2. Count the number of movies released before the 1st of January, 1930

Write a query that counts the number of movies released before the 1st of January, 1930, and run it in the SQL client.

Then, assign the result to variable count_movies_before_1930 (just copy and paste the value).

<!-- grade_id: cell-be58fad2af1735bd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1f2f87262d1b5fc4, points: 1, grade: true, solution: false -->
```python
expected_hash = '3daebbc6dfd81355f1cc9d9565ab4a4a53bda47f6117529409acc7acb55556bb'
assert hashlib.sha256(str(count_movies_before_1930).encode()).hexdigest() == expected_hash
```

### Q3. Count the number of movies with a character named Maria

Write a query that counts the number of movies that have a caracter whose name is "Maria" or "Maria *something else*", and run it in the SQL client.

Then, assign the result to variable count_movies_with_maria (just copy and paste the value).

Hints: check the [LIKE](https://www.postgresql.org/docs/current/static/functions-matching.html#FUNCTIONS-LIKE) and [DISTINCT](https://www.postgresql.org/docs/current/static/sql-select.html#SQL-DISTINCT) keywords for this exercise.

<!-- grade_id: cell-3a7464176782e2cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f1f24a46b0d15dc9, points: 1, grade: true, solution: false -->
```python
expected_hash = '793733573a1dfd14a2e889a11b2ad7b6981de29df813863b528dc1ae99416eeb'
assert hashlib.sha256(str(count_movies_with_maria).encode()).hexdigest() == expected_hash
```

### Q4. Budget of movies with and without Oscars

Write a query that selects the average budget spent in movies that won an oscar, only considering the movies where the budget is higher than 0. The return value should be the number of dollars in millions, rounded to the nearest integer. Run the query in the SQL client.

Assign the result to variable budget_with_oscars (just copy and paste the values).

Then, write a query that selects the average budget spent in movies that didn't win an oscar, only considering the movies where the budget is higher than 0. The return value should be the number of dollars in millions, rounded to the nearest integer. Run the query in the SQL client.

Assign the result to variable budget_without_oscars (just copy and paste the values).

<!-- grade_id: cell-645c2fbe3b16c880, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41d99aabd53c87e7, points: 1, grade: true, solution: false -->
```python
expected_hash = '4523540f1504cd17100c4835e85b7eefd49911580f8efff0599a8f283be6b9e3'
assert hashlib.sha256(str(budget_with_oscars).encode()).hexdigest() == expected_hash

expected_hash_a = '785f3ec7eb32f30b90cd0fcf3657d388b5ff4297f2f9716ff66e9b69c05ddd09'
expected_hash_b = '6f4b6612125fb3a0daecd2799dfd6c9c299424fd920f9b308110a2c1fbd8f443'

assert hashlib.sha256(str(budget_without_oscars).encode()).hexdigest() in [expected_hash_a, expected_hash_b]
```

### Q5. Which is the actor that has appeared the most in Science Fiction movies

What actor appeared the most in movies with the genre "Western"? Run a query in the SQL client to find out the actor's name and assign it to variable most_western_actor.

Also find out the number of times that the actor appeared Western movies and assign that number to variable most_western_actor_movies_count 

In both the cases, just copy and paste the values into the variables.

<!-- grade_id: cell-c5ef88aa51ea99f6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f0f9a337b81de14a, points: 1, grade: true, solution: false -->
```python
expected_hash = '7ddf1ea571c83c4b4839bfef147be4ba603d1d43ec2b02addd7695b12928337a'
assert hashlib.sha256(most_western_actor.encode()).hexdigest() == expected_hash

expected_hash = 'c75cb66ae28d8ebc6eded002c28a8ba0d06d3a78c6b5cbf9b2ade051f0775ac4';
assert hashlib.sha256(str(most_western_actor_movies_count).encode()).hexdigest() == expected_hash
```

### Querying the MoviesDb with pandas

In these exercises, the goal is to query the MoviesDb using pandas.

### Q6. Select the names of the genres without oscars

The connection settings to use in this exercise are the same ones as in the previous exercises.

Write a query that selects the names of the genres that never received an oscar. The result should have a field named "genre", that should be sorted alphabetically.

Then, use pandas to read this query into a DataFrame called df6 with a single column named "genre".

<!-- grade_id: cell-f2cd873c0ac14ca0, points: 0, grade: false, solution: true -->
```python
# Create an engine that allows to to connect to the MoviesDb PostgreSQL database
# engine = sqlalchemy.create_engine(...)
# YOUR CODE HERE
raise NotImplementedError()


# Write the query as specified in the question
# query = ...
# YOUR CODE HERE
raise NotImplementedError()

# Use pandas read_sql_query function to read the query result into a DataFrame
# df6 = pd.read_sql_query(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-28986c4b783959be, points: 1, grade: true, solution: false -->
```python
assert type(engine) == sqlalchemy.engine.base.Engine
assert len(df6) == 5
assert len(df6.columns) == 1
assert df6.columns.tolist()[0] == 'genre'

expected_hash = '271a22f75d73e13fa558d991442ecde8d75a34a0b9f6216d44b116f8612c82dc'
assert hashlib.sha256(df6.iloc[2].genre.encode()).hexdigest() == expected_hash

expected_hash = '28b7b3139218904b246717b0035295e0383450dd57296187b48fe38a7e9ff0ec'
assert hashlib.sha256(df6.iloc[4].genre.encode()).hexdigest() == expected_hash
```

### Q7. Select foreign movies with the highest budgets

In this exercise, we want to query a local SQLite database.
In order to do this, connect to the the_movies.db database, as was done in the learning notebooks.

Hint: the database file we're using is the same as was used in the Learning Notebooks; the file is in the same directory as this Exercises Notebook.

Write a query that selects the imdb_id, original_title, budget and original_language of movies with an original language other than english.
Then, sort the results from the highest to the lowest budget and return only the first 6 rows.

Use pandas to read this query into a DataFrame called df7 with four columns: imdb_id, original_title, budget and original_language.

<!-- grade_id: cell-8d45d74ae0d84110, points: 0, grade: false, solution: true -->
```python
# Create an engine that allows to to connect to the the_movies.db SQLite database
# engine = sqlalchemy.create_engine(...)
# YOUR CODE HERE
raise NotImplementedError()


# Write the query as specified in the question
# query = ...
# YOUR CODE HERE
raise NotImplementedError()

# Use pandas read_sql_query function to read the query result into a DataFrame
# df6 = pd.read_sql_query(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-19abd085a003ab6f, points: 1, grade: true, solution: false -->
```python
assert type(engine) == sqlalchemy.engine.base.Engine
assert len(df7) == 6
assert len(df7.columns) == 4
assert df7.columns.tolist() == ['imdb_id', 'original_title', 'budget', 'original_language']

expected_hash = '2d78c1c57c0cb7aa2098901eb9f55a291ac54206a375d18ce8107fc8f765a4c7'
assert hashlib.sha256(df7.iloc[0].imdb_id.encode()).hexdigest() == expected_hash

expected_hash = '6085aab10f651514a3a04a504b79b1720478630951bd83238f72a36b15a5aa51'
assert hashlib.sha256(df7.iloc[2].imdb_id.encode()).hexdigest() == expected_hash

expected_hash = 'c27f97e4898ee23badcda22a632e4b3ef7d22593ebc7a64e46ac37520843a255'
assert hashlib.sha256(df7.iloc[4].imdb_id.encode()).hexdigest() == expected_hash
```

## Part B - Public APIs

In this exercises, the goal is to get data from a public API. We'll use the [Rick and Morty API](https://rickandmortyapi.com/).

In order to complete the exercises, you'll have to consult the API's [documentation](https://rickandmortyapi.com/documentation).

<br>

<img src="media/rick_and_morty.png" width=600>

<br>

### Q8. Get location 34 from the API

Read the [documentation](https://rickandmortyapi.com/documentation/#get-a-single-location) of the API in order to find out how to get the location with id 34.
In order to get this data, you'll need to do an HTTP GET request.

The result should be a JSON object (which is the same as a dictionary in Python), and assigned to variable location_34.

<!-- grade_id: cell-3e7084c9bf8398d0, points: 0, grade: false, solution: true -->
```python
# Do an HTTP GET request to the Rick and Morty API to get location 34 as a JSON object
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-50d051ce26525944, points: 3, grade: true, solution: false -->
```python
assert type(location_34) == dict
assert set(location_34.keys()) == {'created', 'dimension', 'id', 'name', 'residents', 'type', 'url'}
assert location_34['name'] == "Earth (Evil Rick's Target Dimension)"
assert location_34['type'] == 'Planet'
```

### Q9. Filter characters from the API

Read the documentation of the API in order to find out how to filter characters using a GET request with query parameters.

Then, find the characters with the following properties:
* name: rick
* status: alive
* type: Pickle

Return only the first (and only) result.

The result should be a JSON object (which is the same as a dictionary in Python), that corresponds to the character we're looking for, and assigned to variable mistery_character.

<!-- grade_id: cell-3c68a66b3b6b56d0, points: 0, grade: false, solution: true -->
```python
# Do an HTTP GET request to filter characters according to the criteria above
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-15c07449c3238657, points: 3, grade: true, solution: false -->
```python
assert type(mistery_character) == dict
assert set(mistery_character.keys()) == {'created', 'episode', 'gender', 'id', 'image', 'location', 
                                         'name', 'origin', 'species', 'status', 'type', 'url'}
assert mistery_character['status'] == 'Alive'
assert mistery_character['species'] == 'unknown'
assert mistery_character['type'] == 'Pickle'

expected_hash = '266fb5572dcafb40ba7dae8c4c28f5915221f88e93fa609bcd77b3d76dcd29bc'
assert hashlib.sha256(mistery_character['name'].encode()).hexdigest() == expected_hash
```

## Part C - Web scraping

In this exercise, we're going to use web scraping to get data about Rick and Morty's episodes ratings for the first season.
We're going to read the data from IMDb, in particular from [this web page](https://www.imdb.com/title/tt2861424/episodes?season=1).

### Q10. Scrape Rick and Morty's ratings from IMDb

Assign a list with the ratings values to variable ratings_list.
In the list, each rating should be a float.

<!-- grade_id: cell-bd8d7b19855c5b9f, points: 0, grade: false, solution: true -->
```python
# Assign the URL of the page to be scraped to variable url
# url = ...
# YOUR CODE HERE
raise NotImplementedError()

# Do a GET request to get the page content, using the url we've just defined
# response = ...
# YOUR CODE HERE
raise NotImplementedError()

# Instanciate a soup object using the response of the GET request
# YOUR CODE HERE
raise NotImplementedError()
    
# Now it's the tricky part!
# Parse the soup in order to retrieve the ratings.
# In the end, store the ratings in a list and assign it to variable ratings_list.
# Make sure that all the ratings in the list are floats!
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-335a90721257cd1a, points: 6, grade: true, solution: false -->
```python
assert type(ratings_list) == list
assert len(ratings_list) == 11

assert type(ratings_list[0]) == float
assert ratings_list[0] == 8.0

assert math.isclose(8.5, np.mean(ratings_list), rel_tol=1e-2)
```

### Q11. Where did you find the ratings?

When you were scraping the Rick and Morty's ratings, you found out that the information you needed was in an HTML element, which looks like this:

```
<tagname class=classname>Ratings content goes here</tagname>
```

Regarding the HTML element where you found the rating's value:

* Assign the tagname to variable ratings_tagname
* Assign the classname to variable ratings_classname

In both cases you don't need to code, just copy and paste the values into the two variables.

<!-- grade_id: cell-88e2603722dac1e3, points: 0, grade: false, solution: true -->
```python
# ratings_tagname = ...
# ratings_classname = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f78d31cd4e4779eb, points: 1, grade: true, solution: false -->
```python
expected_hash = '37a0025aa1f2e29f9f88f59f961e44a320849dfc79a85ef3c02654d50d2e7480'
assert hashlib.sha256(ratings_tagname.encode()).hexdigest() == expected_hash

expected_hash = '6a9b4f4e5f862a625f86aa2c0ad343cb5ff23da130b59b0914f36a5af1fbee92'
assert hashlib.sha256(ratings_classname.encode()).hexdigest() == expected_hash
```

```python

```